---
author: Fluent Data, LLC
date: '`r Sys.Date()`'
tutorial:
  id: training.r.intro.1
  version: 1.0
description: This lesson provides a basic introduction to R programming for air quality
  data science, covering installation, basic operations, and data structures.
title: 1 .  Introduction to R
output:
  learnr::tutorial:
    allow_skip: yes
    exercise.reuse: yes
runtime: shiny_prerendered
---


```{r setup, include=FALSE}
 library(learnr)
library(gradethis)
library(trainingRIntro)
library(shiny) 
  
```

This lesson is a part of the Introduction to R for Air Quality Data Science. The sections below provide a basic introduction to R, including how to install and set up R and RStudio, an overview of R syntax, and how to perform simple operations.

## What is R?

R is a free, open-source computing language. It was originally written
by statisticians for doing statistical analysis in academia. In recent
years it has become more widely used in many industries for performing
a variety of data science tasks such as:

- reading and writing files,
- data transformation,
- graphic visualization,
- geographic mapping,
- and predictive modeling.


## Why Use a Programming Language?

R is one of several programming languages that can be used for data science, including [Python](https://www.python.org/) and
[Julia](https://julialang.org/). They each have advantages and disadvantages, but they are all powerful tools for data science.
These [high-level languages](https://en.wikipedia.org/wiki/High-level_programming_language) give you access to modern algorithms
for processing large amounts of data in a few lines of code.


Many data analysis tasks can be accomplished with spreadsheets and other business intelligence (BI) tools such as Looker and Power
BI. When should you move beyond BI tools and use a high-level programming language like R? Below are a few scenarios where a language
like R is more advantageous than a BI tool.

- If you cannot access data easily in your BI tool, R can read just about any data source.
- If you need to download, save, or otherwise process a large number of files, R can automate those tasks in a way that BI tools cannot.
- Custom data transformations that are not possible in BI tools can be done with R.
- Custom data visualizations that are not available in BI tools can be done with R.
- Predictive modeling that is not available in BI tools, or only in a rudimentary way, can be done in R.

BI tools are more advantageous if you need enterprise wide dashboards,
or tools that are more easily accessible to a wider audience. If there
are few occasions where you need custom visualizations or transformations,
or if you do not need automation in your work, you may not need to learn
a programming language.


## Basic Math

Open up a script if you haven't already (<U+201C>File<U+201D> -> <U+201C>New File<U+201D> -> <U+201C>R Script<U+201D>). Try some math by either typing the lines below or copying and pasting
the lines into your script.


```{r ex-81f861c580e9, exercise = FALSE, eval = TRUE, exercise.cap = 'Practice Basic Math Operations'}
10 + 5
10 - 5
10 * 5
10 / 5
10 ^ 5

```

Remember, to run the lines, highlight your code and click the "Run" button on the toolbar of the script panel. Below is a table of the math operators
in the R language.


| Operator | Meaning | Example |
| :--- | :--- | :--- |
| + | addition | 2 + 2 |
| - | subtraction | 2 - 2 |
| * | multiplication | 2 * 2 |
| / | division | 2 / 2 |
| ^ | exponentiation | 2 ^ 2 |

### Order of Operations

R follows the usual order of arithmetical operations and uses parentheses for grouping operations. Run the two lines of code below and you can
see the different values that are returned.


```{r ex-79b4cf434d78, exercise = TRUE, exercise.eval = FALSE, exercise.cap = 'Understanding Order of Operations'}
10 - 3 / 5

```

```{r ex-17bb0057c9a3, exercise = TRUE, exercise.eval = FALSE, exercise.cap = 'Practicing Grouping with Parentheses'}
(10 - 3) / 5

```

## Comments and Code Blocks

To write a comment in your script that will not be evaluated, type `#` in front of your comment. The text after `#` will not be evaluated. There is
no multi-line commenting in R, so every comment line must begin with the `#` character.


Run all of the code below and see what gets returned in the R console (bottom left panel in RStudio).


```{r ex-d3c1f581907f, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Experiment with Comments'}
# Full line comment
5^2 # partial line comment

```

In the example above and the previous section, you can see the R code and its output. The code blocks with output look like this, with `1+1` being
the R code and `## [1] 2` being the output:


```{r ex-ce008c1109f0, exercise = FALSE, eval = TRUE, exercise.cap = 'Code and Output Example'}
1+1

```

However, in the R console the code and output would look like this:


```{r ex-9a36f5215f5f, exercise = FALSE, eval = FALSE, echo = TRUE, exercise.cap = 'Console Code Execution'}
> 1 + 1
[1] 2

```

The R code in code blocks do not show the `>` part of the console, called
the prompt, and the output block places two comment marks `##` before the
output. This is to make it possible to copy and paste the text into your
R console and run it without causing an error.


## Variables

A variable is a letter or combination of alphanumeric characters that is used to store data. To create a variable in R, use the less-than character
with the dash to create an arrow symbol pointing left `<-`. Below, the variables `x` and `y` are created by assigning some numbers to them.


```{r ex-ec1ce06e0ab7, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Creating Variables'}
x <- 10
y <- 5
x + y

```

In RStudio, you will see the variables we created in the top right panel.


<img src='./images/rstudio_variables.png' alt='Variables in RStudio' style='max-width: 100%;' />

If you've already created a variable, you can replace the value with another value.


```{r ex-2df91768bd95, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Exploring Variable Assignment'}
x

```

```{r ex-d0f95f7fc26e, exercise = TRUE, exercise.eval = FALSE, exercise.cap = 'Updating Variable Values'}
x <- 20
x

```

In the top right panel you can see that the number stored in the variable `x` has changed.


<img src='./images/rstudio_updated_variable.png' alt='Updated variable in RStudio' style='max-width: 100%;' />

There are 3 important rules to remember when creating variable names:

1. You can't start your variable name with a number.
2. You can't use spaces or special characters ($,%,#,-). Periods `.` and underscores  `_` are ok.
3. Capitalization __DOES__ matter in R. That is, R will consider `y` and  `Y` to be different variables.


Try running the following code and you will see that in your global environment there are two different objects listed.


```{r ex-14716a5ddf54, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Variable Naming Rules'}
y <- 5
Y <- 10

```

## Data Types

R has three main data types:


| Type | Description | Examples |
| :--- | :--- | :--- |
| character | letters and words | `z`, `red`, `H2O` |
| numeric | numbers | `1`, `3.14`, `log(10)` |
| logical | binary | `TRUE`, `FALSE` |

The `character` type requires single or double quotes. The numeric type
must be unquoted numbers, and the full-caps logical values `TRUE` and
`FALSE` must also be unquoted.


## Grouping Data

There are several ways to store groups of data to make them easier to work with:

- A __vector__ stores multiple values of the same type (e.g. all numeric values).
- A __list__ stores multiple values of different types (e.g. some numbers and character values).
- A __matrix__ is a table of values with only one data type.
- A __data frame__ is a table of values that can have columns with different data types (e.g. a numeric column and a logical column).


### Vectors

A vector variable can contain only one type of data (numeric, character, or logical). We use `c()` to create vectors.


```{r ex-67078c489715, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Creating Numeric Vectors'}
x <- c(1, 2, 3, 4, 5)
x

```

```{r ex-63506426ac71, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Creating Character Vectors'}
fruit <- c("apples", "bananas", "oranges")
fruit

```

If you try to type in text without using quotations marks for character values (either single or double quotes), then R will throw an error. Try
running the code below.


```{r ex-bba8373fea79, exercise = TRUE, exercise.eval = FALSE, exercise.cap = 'Error Without Quotes'}
fruit <- c(apples, bananas, oranges)

```

R will interpret the text without quotes as the names of variables. Since we don't have any variables named `apples`, `bananas`, or `oranges`, R can't
find them and it returns an error. The members of a vector can be accessed by using single brackets `[ ]`. For example, to print the 3rd member of the
`fruit` vector, you can use the single bracket with the number 3:


```{r ex-5cd4385f6056, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Accessing Vector Elements'}
fruit[3]

```

### Lists

Lists are like vectors but can contain any mix of data types. We use `list()` to create a list variable.


```{r ex-6a0e866a73e4, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Creating Lists'}
x <- list("Benzene", 1.3, TRUE)
x

```

When a list is printed to the console, you will see double brackets `[[ ]]` with the index number inside. In the example above, the numeric value `1.3`
is the second value in the list, so it is shown below the double bracket `[[2]]`. You can access the second value by using the double brackets on the
list.


```{r ex-73b1c6677a00, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Accessing List Elements'}
x[[2]]

```

Lists can also contain vectors and other lists.


```{r ex-c3a1b6cfb9ea, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Lists Containing Vectors and Lists'}
my_vector <- c(1, 2, 3)
my_list <- list("Benzene", 1.3, TRUE)
y <- list(TRUE, my_vector, my_list)
y

```

In this example, you can use two double brackets to access the value `1.3` by selecting the 3rd member of the list `y` and the 2nd member of the list
`my_list`:


```{r ex-a4653820754f, exercise = FALSE, exercise.eval = TRUE, exercise.cap = 'Nested List Access'}
y[[3]][[2]]

```

### Data Frames

Data frames are data tables in R. We use `data.frame()` to create a data frame of vectors of the same length. In the example below, we create three
vectors of the same length and use them to create a data frame.


```{r ex-85321c4c22ee, exercise = TRUE, exercise.eval = FALSE, exercise.cap = 'Creating Data Frames'}
pollutant <- c("Benzene", "Toluene", "Xylenes")
concentration <- c(1.3, 5.5, 6.0)
carcinogen <- c(TRUE, FALSE, FALSE)
my_data_frame <- data.frame(pollutant, concentration, carcinogen)
my_data_frame

```

The output above shows a table with the vector variable names as column names, and the vector values below the respective column name. If you try to
create a data frame where the vectors are not all the same length, you will get an error.


```{r ex-2a5a8901cad3, exercise = TRUE, exercise.eval = FALSE, exercise.cap = 'Data Frame Dimension Mismatch Error'}
pollutant <- c("Benzene", "Toluene")
concentration <- c(1.3, 5.5, 6.0)
carcinogen <- c(TRUE, FALSE, FALSE)
my_data_frame <- data.frame(pollutant, concentration, carcinogen)

```

## Further Reading

* For a comprehensive introduction to data science using R, see
[R For Data Science, 2nd Edition](https://r4ds.hadley.nz/).
* For a good introduction to R programming, see
[Hands on R Programming](https://rstudio-education.github.io/hopr/).



## Exercises {data-progressive=TRUE}

### Exercise 1

Add 25 and 75 in R.

```{r exercise1, exercise = TRUE}
# Your code here
```

```{r exercise1-hint-1}
# You can add two numbers directly using the `+` operator.
```

```{r exercise1-hint-2}
# Simply type `25 + 75` in the R console to perform the addition.
```

```{r exercise1-solution}
25 + 75

```

```{r exercise1-check}
grade_this_code(
  correct = c(gradethis::random_praise(), "R allows you to do simple arithmetic operations such as addition, subtraction, multiplication, etc. directly. Here, we're simply adding two numbers to get their sum. ")
)
```


### Exercise 2

Create an R variable `x` that stores the value 10.

```{r exercise2, exercise = TRUE}
# Your code here
```

```{r exercise2-hint-1}
# Use the assignment operator `<-` to assign a value to a variable.
```

```{r exercise2-hint-2}
# To create a variable `x` with the value 10, use `x <- 10`.
```

```{r exercise2-solution}
x <- 10

```

```{r exercise2-check}
grade_this_code(
  correct = c(gradethis::random_praise(), "The arrow symbol (`<-`) is used to assign a value to a variable in R. Here, we assign the value 10 to the variable `x`. ")
)
```


### Exercise 3

Create a vector `v` that contains the numbers 1 through 5.

```{r exercise3, exercise = TRUE}
# Your code here
```

```{r exercise3-hint-1}
# Use the `c()` function to combine multiple values into a vector.
```

```{r exercise3-hint-2}
# You can create a vector of consecutive numbers using `c(1, 2, 3, 4, 5)` or the shorthand `1:5`.
```

```{r exercise3-solution}
v <- c(1, 2, 3, 4, 5)

```

```{r exercise3-check}
grade_this_code(
  correct = c(gradethis::random_praise(), "We use the `c()` function in R to combine elements into a vector. To create consecutive integers, you can also use the colon `:` operator, which is a convenient way to create sequences of numbers. ")
)
```


### Exercise 4

Create a list `l` that contains the number 5, the string 'apple', and the logical value `TRUE`.

```{r exercise4, exercise = TRUE}
# Your code here
```

```{r exercise4-hint-1}
# Use the `list()` function to create a list that can contain elements of different types.
```

```{r exercise4-hint-2}
# To create a list with a number, a string, and a logical value, you can use `list(5, 'apple', TRUE)`.
```

```{r exercise4-solution}
l <- list(5, 'apple', TRUE)

```

```{r exercise4-check}
grade_this_code(
  correct = c(gradethis::random_praise(), "In R, we create a list using the `list()` function. A list can contain elements of different types, such as numbers, strings, and logical values, allowing for flexible data structures. ")
)
```


### Exercise 5

Create a data frame `df` with two columns, `name` and `age`. Use the names 'Alice', 'Bob', and 'Charlie' and give them each an age.

```{r exercise5, exercise = TRUE}
# Your code here
```

```{r exercise5-hint-1}
# Use the `data.frame()` function to create a data frame.
```

```{r exercise5-hint-2}
# To create a data frame with `name` and `age` columns, define each column as a vector and use them as arguments in the `data.frame()` function.
```

```{r exercise5-solution}
df <- data.frame(name = c('Alice', 'Bob', 'Charlie'), age = c(25, 32, 28))

```

```{r exercise5-check}
grade_this({
  # The student's solution
  df_student <- .result

  # Expected names
  expected_names <- c('Alice', 'Bob', 'Charlie')

  # Check that the data frame has the correct structure
  has_correct_columns <- all(names(df_student) == c('name', 'age'))
  contains_all_names <- all(expected_names %in% df_student$name)
  are_ages_numeric <- is.numeric(df_student$age)

  # Provide detailed feedback based on specific issues
  if (!has_correct_columns) {
    fail("Incorrect columns. Ensure your data frame has exactly two columns named `name` and `age`.")
  } else if (!contains_all_names) {
    missing_names <- setdiff(expected_names, df_student$name)
    fail(paste("Missing the following names:", paste(missing_names, collapse = ", ")))
  } else if (!are_ages_numeric) {
    fail("The `age` column should contain numeric values only. Ensure all ages are numbers.")
  } else {
    pass("Correct! You've successfully created a data frame with the names 'Alice', 'Bob', and 'Charlie', and ensured that ages are numeric.")
  }
})

```


### Exercise 6

Create a numeric vector `nums` with the values 10, 20, and 30 and add 5 to each element of the vector.

```{r exercise6, exercise = TRUE}
# Your code here
```

```{r exercise6-hint-1}
# You can create a vector using the `c()` function.
```

```{r exercise6-hint-2}
# You can create a vector using the `c()` function.
# `nums <- c(10, 20, 30)`
```

```{r exercise6-hint-3}
# To add a number to each element of a vector, you can simply use the `+` operator with the vector and the number.
```

```{r exercise6-hint-4}
# To add a number to each element of a vector, you can simply use the `+` operator with the vector and the number.
# `nums + 5`
```

```{r exercise6-solution}
nums <- c(10, 20, 30)
nums + 5
```

```{r exercise6-check}
grade_this_code(
  correct = c(gradethis::random_praise(), "We can operate on every element of a vector at once in R. By creating a vector `nums` with `c(10, 20, 30)` and then adding 5 to it, we perform an element-wise addition, resulting in each element of the vector being increased by 5. ")
)
```




## Next Lesson

You have completed Lesson . Click the button below to mark it as complete and move on to the next lesson.

```{r example-button, echo=FALSE}
  actionButton("complete_lesson", "Mark Lesson 1 Complete")
```


```{r, context = "server"}
observeEvent(input$complete_lesson, {
  showModal(modalDialog(
    title = 'Congratulations!',
tags$p("Congratulations on completing Introduction to R! You've mastered the essentials of R setup and basic operations. Next up, Functions and Importing Data, where you'll learn to enhance your data manipulation skills with R functions and importing techniques.
"),
    easyClose = FALSE,
    footer = tagList(
      modalButton('Cancel'),
      actionButton('confirm_complete', 'Confirm', class = 'btn-primary')
    )
  ))
})

observeEvent(input$confirm_complete, {
  removeModal()
  trainingRIntro::set_user_state(lesson_1_complete=TRUE)
  shiny::stopApp()
})



